```python
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
warnings.simplefilter(action='ignore', category=DeprecationWarning)

import numpy as np
import pandas as pd
import seaborn as sns
```

The dataset `Aerial_Waterfowl_Survey_Data.csv` shows counts of various species of waterfowl observed during aerial surveys:

```python
birds = pd.read_csv("Aerial_Waterfowl_Survey_Data.csv")
birds
```

## 1.
(2.1) Make a data frame called `ducks` for just the species *American Black Duck*, *Mallard*, *Ruddy Duck*, and *Ring-necked Duck*. Make a series called `duckmeans` for the means of those 4 columns.

<!-- grade_id: ducks-mean, points: 0, grade: false, solution: true -->
```python
ducks = None
duckmeans = None

# YOUR CODE HERE
raise NotImplementedError()

print(ducks.head(5))
print()
print(duckmeans)
```

<!-- grade_id: ducks-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert np.isclose(duckmeans.sum(), 1834.528606356), "The mean values are incorrect"
assert len(ducks.columns) == 4, "Wrong number of columns in ducks"
print("OK")
```

<!-- grade_id: ducks-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 2.
(2.1) For the *Canada Goose* column of `birds`, assign `CGmu` the value of the mean, assign `CGsigma` the value of the standard deviation, and let `CGZ` be a series of the z-scores of the column.

<!-- grade_id: stats-answer, points: 0, grade: false, solution: true -->
```python
CGmu = None
CGsigma = None
CGZ = None

# YOUR CODE HERE
raise NotImplementedError()

print(f"mean for Canada Goose: {CGmu:.2f}")
print(f"std for Canada Goose: {CGsigma:.2f}")
```

<!-- grade_id: cell-fb7bce0cda1d50ca, points: 2, grade: true, solution: false -->
```python
# TESTS
assert np.isclose(CGZ.mean(), 0), "The mean of CGZ should be near zero"
assert np.isclose(CGZ.std(), 1), "The standard deviation of CGZ should be near 1"
assert len(CGZ) == birds.shape[0], "CGZ should have the same number of rows as birds"
assert np.isclose(CGZ.median(), -0.34007213204), "The CGZ values are incorrect"
print("OK")
```

<!-- grade_id: stats-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 3.
(2.1) Let `IQR` be a series of the interquartile ranges of all the columns of *birds* after the 5th, keeping only the values greater than zero.

<!-- grade_id: iqr-answer, points: 0, grade: false, solution: true -->
```python
IQR = None

# YOUR CODE HERE
raise NotImplementedError()
IQR
```

<!-- grade_id: iqr-tests, points: 0, grade: true, solution: false -->
```python
# TESTS
assert type(IQR) == pd.Series, "IQR should be a pandas Series"
assert np.isclose(IQR.sum(), 14026), "The sum of IQR values is incorrect"
assert np.isclose(IQR["Snow Goose"], 5900), "The Snow Goose IQR value is incorrect"
print("OK")
```

<!-- grade_id: iqr-grade, points: 1, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 4.
(2.2) Plot the empirical CDF of the *Snow Goose* column using only the rows in which *Zone* equals 3.

<!-- grade_id: ecdf-manual, points: 3, grade: true, solution: true -->
```python
# Result is a plot
# YOUR CODE HERE
raise NotImplementedError()
```

## 5.
(2.2) Plot a histogram with 25 bins for the **positive entries only** in the *Bufflehead* column.

<!-- grade_id: hist-manual, points: 3, grade: true, solution: true -->
```python
# Result is a plot
# YOUR CODE HERE
raise NotImplementedError()
```

## 6.
(2.3) Make a series called `canada` of the total number of *Canada Goose* observed in each zone. I.e., the index has the zone numbers, and the values are the total number in each zone.

<!-- grade_id: canada-answer, points: 0, grade: false, solution: true -->
```python
canada = None
# YOUR CODE HERE
raise NotImplementedError()

print(canada)
```

<!-- grade_id: canada-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(canada) == pd.Series, "canada should be a pandas Series"
assert len(canada) == 11, "canada should have 11 values"
assert canada[3.0] == 1968090, "The value for Zone 3 is incorrect"
print("OK")
```

<!-- grade_id: canada-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 7.

(2.3) Make a box plot, grouped by month, of the *Snow Goose* counts for all the years since 2000 (inclusive).

<!-- grade_id: mallard-plot, points: 4, grade: true, solution: true -->
```python
# Result is a plot
# YOUR CODE HERE
raise NotImplementedError()
```

## 8.
(2.4) Produce code to count the total number of outliers for December in the previous plot and assign to `numOutliersDec`.

<!-- grade_id: outliers-answer, points: 0, grade: false, solution: true -->
```python
numOutliersDec = None
# YOUR CODE HERE
raise NotImplementedError()

print(numOutliersDec)
```

<!-- grade_id: outliers-test, points: 0, grade: true, solution: false -->
```python
# TESTS
assert numOutliersDec > 10, "There should be more than 10 outliers in December"
assert numOutliersDec < 25, "There should be fewer than 25 outliers in December"
print("OK")
```

<!-- grade_id: outliers-grade, points: 4, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 9. 
(2.5) Create a 4x4 grid of pairwise scatter plots for the `ducks` frame from #1.

<!-- grade_id: ducks-plot, points: 2, grade: true, solution: true -->
```python
# Results in a plot

# YOUR CODE HERE
raise NotImplementedError()
```

## 10.
(2.5) Assign to `duckduck` a dataframe of the Pearson correlation coefficients between all pairings of the four species in `ducks`.

<!-- grade_id: duckduck-answer, points: 0, grade: false, solution: true -->
```python
duckduck = None
# YOUR CODE HERE
raise NotImplementedError()

duckduck
```

<!-- grade_id: duckduck-tests, points: 1, grade: true, solution: false -->
```python
# TESTS
assert type(duckduck) == pd.DataFrame, "duckduck should be a DataFrame"
assert duckduck.shape == (4, 4), "duckduck should be a 4x4 DataFrame"
assert np.isclose(duckduck.loc['Mallard', 'Ruddy Duck'], -0.027620580191), "A correlation value is incorrect"
print("OK")
```

<!-- grade_id: duckduck-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 11.
(2.5)
Repeat the previous problem, but using Spearman coefficients, and call the result `spearduck`.

<!-- grade_id: spearduck-answer, points: 0, grade: false, solution: true -->
```python
spearduck = None
# YOUR CODE HERE
raise NotImplementedError()

spearduck
```

<!-- grade_id: spearduck-tests, points: 1, grade: true, solution: false -->
```python
# TESTS
assert type(spearduck) == pd.DataFrame, "spearduck should be a DataFrame"
assert np.isclose(spearduck.loc['Mallard', 'Ruddy Duck'], 0.0634493133516), "A correlation value is incorrect"
print("OK")
```

<!-- grade_id: spearduck-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```