```python
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
warnings.simplefilter(action='ignore', category=DeprecationWarning)

import pandas as pd
import seaborn as sns
import numpy as np
```

This dataset describes arrests for serious offenses in Delaware from 2002 to 2020:

```python
crimes = pd.read_csv("Delaware_Crime_Report_Totals.csv")
crimes
```

# 1.
(1.6) Let `beaches` be the data frame consisting of the rows from `crimes` where the *Jurisdiction* column is one of "Bethany Beach", "Dewey Beach", or "Rehoboth Beach". (Hint: The `isin` method is quite useful here.)

<!-- grade_id: beaches-answer, points: 0, grade: false, solution: true -->
```python
beaches = None
# YOUR CODE HERE
raise NotImplementedError()
beaches
```

<!-- grade_id: beaches-test, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(beaches) == pd.DataFrame, "beaches should be a DataFrame"
assert len(beaches) == 3318, "beaches should have 3318 rows"
assert np.isclose(beaches["Clearance Rate"].sum(), 72835.75), "The sum of the Clearance Rate column is incorrect"
print("OK")
```

<!-- grade_id: beaches-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

# 2.
(2.3) Make a scatterplot of `beaches` showing *Year* on the x axis and *Adult Arrests* on the y axis, using color to indicate the value of *Offense Type* and facet column to indicate the value of *Jurisdiction*.

<!-- grade_id: scatter-manual, points: 0, grade: true, solution: true -->
```python
# Result is a plot
# YOUR CODE HERE
raise NotImplementedError()
```

# 3.
(2.3) Make a series indexed by *Year* showing the total number of *Adult Arrests* in which *Main Offense* is "Prostitution".

<!-- grade_id: prost-answer, points: 0, grade: false, solution: true -->
```python
prost = None
# YOUR CODE HERE
raise NotImplementedError()
print(prost)
```

<!-- grade_id: prost-test, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(prost) == pd.Series, "prost should be a Series"
assert len(prost) == 19, "prost should have 19 elements"
assert sum(prost) == 7403, "The sum of prost is incorrect"
print("OK")
```

<!-- grade_id: prost-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

# 4.
(2.5) Find the correlations between *Adult Arrests* and *Juvenile Arrests* when `crimes` is grouped by *Offense Type*. Which offense type shows the weakest correlation?

<!-- grade_id: corr-answer, points: 0, grade: false, solution: true -->
```python
correlations = None
weakest = ""

# YOUR CODE HERE
raise NotImplementedError()

print(correlations)
print()
print("Weakest offense type is", weakest)
```

<!-- grade_id: corr-test, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(correlations) == pd.DataFrame, "correlations should be a DataFrame"
assert correlations.shape == (8,2), "correlations should have 8 rows and 2 columns"
assert type(weakest) == str, "weakest should be a string"
assert len(weakest) > 0, "weakest should not be an empty string"
print("OK")
```

<!-- grade_id: corr-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```