```python
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
warnings.simplefilter(action='ignore', category=DeprecationWarning)

import pandas as pd
import seaborn as sns
import numpy as np
```

This dataset contains estimates and projections of greenhouse gas emissions in Delaware:

```python
emissions = pd.read_csv("Delaware_Greenhouse_Gas_Emissions.csv").dropna()
amount = 'MMTCO2e'
emissions
```

## 1.
(1.6) Let `co2` be the frame resulting from selection of the rows in which *Category* is `'Estimate'` and *Gas* is `'CO2'`.

<!-- grade_id: co2-answer, points: 0, grade: false, solution: true -->
```python
co2 = None

# YOUR CODE HERE
raise NotImplementedError()

co2
```

<!-- grade_id: co2-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(co2) == pd.DataFrame, "co2 should be a DataFrame"
assert co2.shape == (406,8), "co2 should be 406 by 8"
assert all( co2["Category"]=="Estimate" ), "All rows in co2 should have Category=Estimate"
assert all( co2["Gas"]=="CO2" ), "All rows in co2 should have Gas=CO2"
print("OK")
```

<!-- grade_id: co2-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 2.
(2.3) Make a box plot of *MMTCO2e* from `co2`, putting *Economic sectors* on the y axis and separating facet columns according to *Fuel type*.

<!-- grade_id: boxplot-answer, points: 4, grade: true, solution: true -->
```python
# Result is a plot
# YOUR CODE HERE
raise NotImplementedError()
```

## 3.
(2.5) Make line plots from `co2` with *Year* on the x axis, *MMTCO2e* on the y axis, colors for *Economic sectors*, and columns for *Fuel type*. (There will be no error bars/ribbons on these.)

<!-- grade_id: lineplot-answer, points: 4, grade: true, solution: true -->
```python
# Result is a plot
# YOUR CODE HERE
raise NotImplementedError()
```

## 4.
(2.6) Here is the correlation between *Year* and emissions of CO2 from `co2`:

```python
co2[["Year", "MMTCO2e"]].corr()
```

Let `sectorCorr` be the same result but with the correlations grouped by *Economic sectors*. 

<!-- grade_id: corr-answer, points: 0, grade: false, solution: true -->
```python
sectorCorr = None
# YOUR CODE HERE
raise NotImplementedError()
sectorCorr
```

<!-- grade_id: cell-d840a8053b5fc5b8, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(sectorCorr) == pd.DataFrame, "sectorCorr should be a DataFrame"
assert sectorCorr.shape == (10,2), "sectorCorr should be 10 by 2"
assert sectorCorr.loc[("Commercial","Year"),"MMTCO2e"] > 0, "Correlation between Year and MMTCO2e for Commercial should be positive"
assert sectorCorr.loc[("Electric power","Year"),"MMTCO2e"] < 0, "Correlation between Year and MMTCO2e for Electric power should be negative"
print("OK")
```

<!-- grade_id: corr-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```