```python
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
warnings.simplefilter(action='ignore', category=DeprecationWarning)

import numpy as np
import pandas as pd
import seaborn as sns
```

This dataset is a collection of the system permissions of Android apps, with the *Result* column indicating whether the app is malware.

```python
perms = pd.read_csv("android_permissions.csv")
perms.head()
```

## 1.
(3.1)
Let `X` be a feature frame found by dropping the *Result* column, and let `y` be the *Result* column.

<!-- grade_id: Xy-answer, points: 0, grade: false, solution: true -->
```python
X, y = None, None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
X.head()
```

```python
y.head()
```

<!-- grade_id: Xy-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert X.shape == (29332, 86), "X has the wrong shape"
assert y.sum() == 14700, "y has the wrong values"
assert X["android.permission.BLUETOOTH"].sum() == 1219, "Something is wrong with the data in X"
print("OK")
```

## 2. 

(3.2)
Split the dataset, reserving 20% of it as `X_test,y_test` for testing and the rest as `X_train,y_train` for training. 

**IMPORTANT**: Make sure the split order is randomized starting from random state 3383.

<!-- grade_id: split-answer, points: 0, grade: false, solution: true -->
```python
X_train, y_train = None, None
X_test, y_test = None, None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
X_test.head()
```

<!-- grade_id: split-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert len(X_train) == 23465, "X_train has the wrong number of rows"
assert X_test.iloc[:, 10].sum() == 67, "X_test has the wrong rows or columns"
assert y_train.sum() == 11749, "y_train has the wrong rows"
print("OK")
```

## 3.

(3.3) Train a decision tree classifier with maximum depth 12 on the training set. **Make sure to set its random state equal to 302.** Find the accuracy of the decision tree on the test set.

<!-- grade_id: acc-answer, points: 0, grade: false, solution: true -->
```python
dtree, acc = None, None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f"Decision tree accuracy is {acc:.3%}")
```

<!-- grade_id: acc-test, points: 2, grade: true, solution: false -->
```python
# TESTS
assert 0.94 < acc < 0.97, "Incorrect accuracy score"
print("OK")
```

<!-- grade_id: acc-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally left blank--do not delete
```

## 4.

(3.5) Find the probability, according to the decision tree, that the app in the first row of `X_test` is malware.

<!-- grade_id: mal-answer, points: 0, grade: false, solution: true -->
```python
mal = None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f"Probability of malware is {mal:.2%}")
```

<!-- grade_id: mal-test, points: 1, grade: true, solution: false -->
```python
# TESTS
assert mal > 0.93, "Prediction is incorrect"
print("OK")
```

<!-- grade_id: mal-grade, points: 3, grade: true, solution: false -->
```python
# Intentionally left blank--do not delete
```

## 5.

(3.5) Plot the ROC curve over the test set for predicting malware (i.e, when *Result* equals 1).

<!-- grade_id: roc-manual, points: 4, grade: true, solution: true -->
```python
# Result is a plot

# YOUR CODE HERE
raise NotImplementedError()
```

## 6. 

(3.5) Find the AUC score for the ROC curve.

<!-- grade_id: auc-answer, points: 0, grade: false, solution: true -->
```python
AUC = None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f"AUC score is {AUC:.5f}")
```

<!-- grade_id: auc-test, points: 1, grade: true, solution: false -->
```python
# TESTS
assert 1 > AUC > 0.95, "AUC is incorrect"
print("OK")
```

<!-- grade_id: auc-grade, points: 1, grade: true, solution: false -->
```python
# Intentionally left blank--do not delete
```