```python
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
warnings.simplefilter(action='ignore', category=DeprecationWarning)

import numpy as np
import pandas as pd
import seaborn as sns # type: ignore
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[1], line 7
      5 import numpy as np
      6 import pandas as pd
----> 7 import seaborn as sns # type: ignore

ModuleNotFoundError: No module named 'seaborn'
```

In this lab you will use data from a survey done during the swine flu (H1N1) scare of 2009-2010 to see how demographcs and opinions affected the likelihood of getting the H1N1 vaccine.

## 1.
Import the `h1n1.csv` file. Drop all rows that have a missing value. Let `y` be the column *h1n1_vaccine*, and let `X` be the result of conversion to dummy variables of all the other columns. Use `drop_first=True` in the dummy conversion.

<!-- grade_id: import-answer, points: 0, grade: false, solution: true -->
```python
X = None
y = None

df = pd.read_csv("h1n1.csv")

df = df.dropna()

y = df['h1n1_vaccine']

X = pd.get_dummies(df.drop(columns=['h1n1_vaccine']), drop_first=True)
```

```python
X
```

Output:
```
h1n1_concern  h1n1_knowledge  behavioral_antiviral_meds  \
0               1.0             0.0                        0.0   
1               3.0             2.0                        0.0   
3               1.0             1.0                        0.0   
4               2.0             1.0                        0.0   
5               3.0             1.0                        0.0   
...             ...             ...                        ...   
26700           3.0             1.0                        0.0   
26701           2.0             2.0                        0.0   
26702           2.0             0.0                        0.0   
26703           1.0             2.0                        0.0   
26706           0.0             0.0                        0.0   

       behavioral_avoidance  behavioral_face_mask  behavioral_wash_hands  \
0                       0.0                   0.0                    0.0   
1                       1.0                   0.0             
```

<!-- grade_id: import-tests, points: 4, grade: true, solution: false -->
```python
# TESTS
assert X.shape == (19682, 37), "X has the wrong shape"
assert y.sum() == 4521, "y has the wrong values"
print("OK")
```

Output:
```
OK
```

## 2.
(4.1) The code below defines a function `learn_curve(model, X, y)` that plots learning curves of precision for a given classifier and dataset. Use it to make learning curves for two decision trees, one with maximum depth 5 and the other with maximum depth 9.

The differences between the plots show clearly that the deeper tree is overfitting; i.e., more data would be needed to make the training match the testing.

```python
# Do not edit this cell
from sklearn.model_selection import LearningCurveDisplay, StratifiedShuffleSplit
cv = StratifiedShuffleSplit(n_splits=20, test_size=0.2, random_state=302)
def learn_curve(model, X, y):
    return LearningCurveDisplay.from_estimator(
    model, X, y,
    random_state=302,
    train_sizes=np.linspace(0.1, 1, 12),
    cv=cv,
    score_name="Precision",
    scoring="precision",
    score_type="both",    # training and testing scores
    n_jobs=-1             # use all processors
    )

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[8], line 2
      1 # Do not edit this cell
----> 2 from sklearn.model_selection import LearningCurveDisplay, StratifiedShuffleSplit
      3 cv = StratifiedShuffleSplit(n_splits=20, test_size=0.2, random_state=302)
      4 def learn_curve(model, X, y):

ModuleNotFoundError: No module named 'sklearn'
```

<!-- grade_id: lc1-answer, points: 2, grade: true, solution: true -->
```python
from sklearn.tree
import DecisionTreeClassifier
import matplotlib.pyplot as plt
model_depth_5 = DecisionTreeClassifier(max_depth=5, random_state=302)
learn_curve(model_depth_5, X, y)
plt.title("Learning Curve - Decision Tree (max_depth=5)")
plt.show()
```

Output:
```
Cell In[9], line 1
    from sklearn.tree
                     ^
SyntaxError: invalid syntax
```

<!-- grade_id: lc2-answer, points: 2, grade: true, solution: true -->
```python
from sklearn.tree import DecisionTreeClassifier
import matplotlib.pyplot as plt
model_depth_9 = DecisionTreeClassifier(max_depth=9, random_state=302)
learn_curve(model_depth_9, X, y)
plt.title("Learning Curve - Decision Tree (max_depth=9)")
plt.show()
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[10], line 1
----> 1 from sklearn.tree import DecisionTreeClassifier
      2 import matplotlib.pyplot as plt
      3 model_depth_9 = DecisionTreeClassifier(max_depth=9, random_state=302)

ModuleNotFoundError: No module named 'sklearn'
```

## 3.

(4.1) Split the dataset 80% / 20% into training and test subsets, using `19716` as the starting random state. Let `tree` be a decision tree classifier with maximum depth 9 (the overfit situation above) and train it using the training subset. Finally, compute its balanced accuracy score on the test set.

<!-- grade_id: tree-answer, points: 0, grade: false, solution: true -->
```python
tree = None
tree_score = None

from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
from sklearn.metrics import balanced_accuracy_score

# Split the dataset
X_train, X_test, y_train, y_test = train_test_split(
    X, y, test_size=0.2, random_state=19716, stratify=y)

# Train the decision tree with max_depth=9
tree = DecisionTreeClassifier(max_depth=9, random_state=19716)
tree.fit(X_train, y_train)

# Predict and evaluate
y_pred = tree.predict(X_test)
score = balanced_accuracy_score(y_test, y_pred)

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[13], line 4
      1 tree = None
      2 tree_score = None
----> 4 from sklearn.model_selection import train_test_split
      5 from sklearn.tree import DecisionTreeClassifier
      6 from sklearn.metrics import balanced_accuracy_score

ModuleNotFoundError: No module named 'sklearn'
```

```python
print(f"Tree balanced accuracy is {tree_score:.2%}")
```

<!-- grade_id: tree-tests, points: 3, grade: true, solution: false -->
```python
# TESTS
assert type(tree) == DecisionTreeClassifier, "tree must be a DT classifier"
assert tree.n_classes_ == 2, "tree must be trained to 2 classes"
assert 0.65 < tree_score < 0.7, "incorrect score for the tree"
print("OK")

```

## 4.

(4.3) Repeat step 3 for a random forest consisting of 100 trees with maximum depth 9 and maximum samples and features set to 50%. Use `3383` as the random state for the forest.

<!-- grade_id: forest-answer, points: 0, grade: false, solution: true -->
```python
forest = None
forest_score = None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f"Forest balanced accuracy is {forest_score:.2%}")
```

<!-- grade_id: forest-tests, points: 3, grade: true, solution: false -->
```python
# TESTS
assert type(forest) == RandomForestClassifier, "forest must be a RF classifier"
assert forest.n_classes_ == 2, "forest must be trained to 2 classes"
assert 0.7 < forest_score < 0.75, "incorrect score for the forest"
print("OK")

```

## 5.

(4.3) Make ROC plots for the `tree` and `forest` classifiers on the test data.

The forest's curve should be better than the tree's.

<!-- grade_id: roc1-answer, points: 2, grade: true, solution: true -->
```python
# PLOT 1
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: roc2-answer, points: 2, grade: true, solution: true -->
```python
# PLOT 2
# YOUR CODE HERE
raise NotImplementedError()
```

## 6.
(4.4) Using a stratified k-fold with 5 splits and random state `302`, find cross-validation scores for the `forest` using all the data (not just test or train). They should be fairly uniform.

<!-- grade_id: cv-answer, points: 0, grade: false, solution: true -->
```python
scores = None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("Validation scores:")
print( scores["test_score"] )
```

<!-- grade_id: cv-tests, points: 4, grade: true, solution: false -->
```python
# TESTS
assert type(scores) == dict, "Result should be a dictionary"
assert (0.7 < scores["test_score"]).all(), "Some scores are too small"
assert (0.75 > scores["test_score"]).all(), "Some scores are too big"
print("OK")

```