```python
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
warnings.simplefilter(action='ignore', category=DeprecationWarning)

import numpy as np
import pandas as pd
import seaborn as sns
from sklearn.neighbors import KNeighborsClassifier
from sklearn.preprocessing import StandardScaler
from sklearn.pipeline import Pipeline
```

Pulsars are detected by scanning for radio frequencies from deep space. However, the vast majority of collected signals are actually terrestrial noise. The following dataset collects 8 statistical features from radio signals and their classifications as noise (class 0) or pulsar (class 1).

```python
signals = pd.read_csv("pulsars.csv")
signals

```

## 1.

Let `y` be the `class` column of the dataset and let `X` be all the remaining columns. Find the proportion of the samples that are actually pulsars.

<!-- grade_id: import-answer, points: 0, grade: false, solution: true -->
```python
X, y = None, None
pulsar_fraction = None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
X
```

```python
print(f"{pulsar_fraction:.1%} of the samples are pulsars.")
```

<!-- grade_id: import-tests, points: 0, grade: true, solution: false -->
```python
# TESTS
assert X.shape == (17898, 8)
assert y.shape == (17898,)
assert 0 < pulsar_fraction < 1, "Pulsar fraction value is wrong"
print("OK")

```

## 2.

Using a random state `19716`, split the data 80% / 20% into testing and training sets.

<!-- grade_id: split-answer, points: 0, grade: false, solution: true -->
```python
X_train, X_test = None, None
y_train, y_test = None, None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
X_train
```

<!-- grade_id: split-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert X_train.shape == (14318, 8)
assert y_test.sum() == 326
```

## 3.
Fit a pipeline with standardization scaling and a kNN classifier with $k=8$ to the training data. Since we want to minimize false positives, find its precision score on the test set.

<!-- grade_id: pipe-answer, points: 0, grade: false, solution: true -->
```python
knn_score = None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f"Test score: {knn_score:.2%}")
```

<!-- grade_id: pipe-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert 0.9 < knn_score < 0.94
print("OK")

```

## 4.

(4.4) Perform a grid search on the pipeline defined in step 3, including a search over $k$ from 3 through 20 and over `'uniform'` and `'distance'` for `weights`, using 6 folds in cross-validation and precision as the scoring. Find the best parameters and find the precision score on the test set.

<!-- grade_id: grid-answer, points: 0, grade: false, solution: true -->
```python
best_params = None
grid_score = None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f"Best parameters: {best_params}")
print(f"Test score: {grid_score:.2%}")
```

<!-- grade_id: grid-tests, points: 5, grade: true, solution: false -->
```python
# TESTS
assert type(best_params) == dict, "Get the best parameters from the fitted model"
assert grid_score > knn_score, "Score should have improved"
print("OK")

```

## 5.

(4.3) Using the best model from step 4, make a bagging classifier using 200 estimators, 50% max features and max samples, and random state `302`. Find its precision score.

<!-- grade_id: bag-answer, points: 0, grade: false, solution: true -->
```python
ensemble_score = None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(f"Test score: {ensemble_score:.2%}")
```

<!-- grade_id: bag-tests, points: 5, grade: true, solution: false -->
```python
# TESTS
assert ensemble_score > grid_score, "Score should have improved"
print("OK")

```