```python
import numpy as np
import pandas as pd
```

## 1.
Create a series called `schools` that has the values `"Delaware"`, `"Penn"`, `"Drexel"`, and `"Villanova"`, in that order.

<!-- grade_id: schools-answer, points: 0, grade: false, solution: true -->
```python
import pandas as pd
schools = None

schools = pd.Series(["Delaware", "Penn", "Drexel", "Villanova"])

print(schools)
```

Output:
```
0     Delaware
1         Penn
2       Drexel
3    Villanova
dtype: object
```

<!-- grade_id: schools-type, points: 1, grade: true, solution: false -->
```python
# TESTS
assert type(schools) == pd.Series
print("OK")
```

Output:
```
OK
```

<!-- grade_id: schools-values, points: 3, grade: true, solution: false -->
```python
# TESTS
assert np.all(schools.values == ["Delaware", "Penn", "Rutgers", "Villanova"])
print("OK")
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[11], line 2
      1 # TESTS
----> 2 assert np.all(schools.values == ["Delaware", "Penn", "Rutgers", "Villanova"])
      3 print("OK")

AssertionError:
```

## 2.
Give `schools` the name "School". Then use it to make a data frame called `uni` that looks like this:

![data frame](./df_snapshot.jpg)

Note: The *Founded* values should be of `int64` data type.

<!-- grade_id: uni-answer, points: 0, grade: false, solution: true -->
```python
uni = None

uni = pd.DataFrame({
    "School": ["Delaware", "Penn", "Rutgers", "Villanova"],
    "Founded": [1743, 1740, 1766, 1842],
    "Mascot": ["YouDee", "Quaker", "Sir Henry", "Will D. Cat"]
}).set_index("School")


# Ensuring "Founded" is of int64 type
uni["Founded"] = uni["Founded"].astype("int64")
print(uni)
```

Output:
```
Founded       Mascot
School                         
Delaware      1743       YouDee
Penn          1740       Quaker
Rutgers       1766    Sir Henry
Villanova     1842  Will D. Cat
```

<!-- grade_id: uni-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(uni) == pd.DataFrame
assert list(uni.columns) == ['Founded', 'Mascot']
assert uni.loc["Penn", "Mascot"] == "Quaker"
assert uni.loc["Villanova", "Founded"] == 1842
print("OK")
```

Output:
```
OK
```

<!-- grade_id: uni-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 3.
 Load the `fight-songs.csv` file as a data frame called `songs`. Set the index of the frame to be the value in the *school* column. Display the first 5 rows.

<!-- grade_id: songs-manual, points: 2, grade: true, solution: true -->
```python
songs = None

songs = pd.read_csv("fight-songs.csv", index_col="school")

# Display the first 5 rows
print(songs.head())
```

Output:
```
conference          song_name  \
school                                         
Notre Dame    Independent      Victory March   
Baylor             Big 12          Old Fight   
Iowa State         Big 12  Iowa State Fights   
Kansas             Big 12      I'm a Jayhawk   
Kansas State       Big 12    Wildcat Victory   

                                                        writers  year  \
school                                                                  
Notre Dame                     Michael J. Shea and John F. Shea  1908   
Baylor                               Dick Baker and Frank Boggs  1947   
Iowa State    Jack Barker, Manly Rice, Paul Gnam, Rosalind K...  1930   
Kansas                                    George "Dumpy" Bowles  1912   
Kansas State                                  Harry E. Erickson  1927   

             student_writer official_song contest  bpm  sec_duration fight  \
school                                                                       
Notre Dame
```

<!-- grade_id: songs-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert songs.shape == (65, 22)
assert songs.index.name == "school"
```

## 4.
Assign the *song_name* column from `songs` to a variable called `names`. Define a new series called `subset` having the 8th through 18th values of `names`.

<!-- grade_id: names-answer, points: 0, grade: false, solution: true -->
```python
names, subset = None, None

songs = pd.read_csv("fight-songs.csv", index_col="school")

names = songs["song_name"]

subset = names.iloc[7:18]

print(names.head(6))
print()
print(subset)
```

Output:
```
school
Notre Dame          Victory March
Baylor                  Old Fight
Iowa State      Iowa State Fights
Kansas              I'm a Jayhawk
Kansas State      Wildcat Victory
Oklahoma            Boomer Sooner
Name: song_name, dtype: object

school
Texas                       Texas Fight
TCU                      TCU Fight Song
Texas Tech        Fight, Raiders, Fight
West Virginia        Hail West Virginia
Illinois                  Oskee-Wow-Wow
Indiana            Indiana, Our Indiana
Iowa                    Iowa Fight Song
Maryland            Maryland Fight Song
Michigan                    The Victors
Michigan State          Victory for MSU
Minnesota          The Minnesota Rouser
Name: song_name, dtype: object
```

<!-- grade_id: names-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert type(names) == pd.Series
assert len(names) == 65
assert type(subset) == pd.Series
assert len(subset) == 11
print("OK")
```

Output:
```
OK
```

<!-- grade_id: names-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 5.
Assign the maximum value in the *trope_count* column of `songs` to the variable `trope_max`.

<!-- grade_id: tropemax-answer, points: 0, grade: false, solution: true -->
```python
trope_max = None

songs = pd.read_csv("fight-songs.csv", index_col="school")

trope_max = songs["trope_count"].max()

print(f"Max tropes is {trope_max}")
```

Output:
```
Max tropes is 8
```

<!-- grade_id: tropemax-tests, points: 0, grade: true, solution: false -->
```python
# TESTS
assert type(trope_max) in [np.int64, int]
print("OK")
```

Output:
```
OK
```

<!-- grade_id: tropemax-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 6.
Make a data frame called `songs1914` consisting of only the rows in `songs` representing songs that were written in 1914.

<!-- grade_id: 1915-answer, points: 0, grade: false, solution: true -->
```python
songs1914 = None

songs = pd.read_csv("fight-songs.csv", index_col="school")

songs1914 = songs[songs["year"] == 1914]

print(songs1914[["conference", "song_name", "year"]])
```

Output:
```
Empty DataFrame
Columns: [conference, song_name, year]
Index: []
```

<!-- grade_id: 1915-tests, points: 1, grade: true, solution: false -->
```python
# TESTS
assert type(songs1914) == pd.DataFrame
assert set(songs1914["year"]) == {"1914"}
print("OK")
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[28], line 3
      1 # TESTS
      2 assert type(songs1914) == pd.DataFrame
----> 3 assert set(songs1914["year"]) == {"1914"}
      4 print("OK")

AssertionError:
```

<!-- grade_id: 1915-grade, points: 3, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 7.
Assign to `fast_corny` the subset of `songs` consisting of the columns *song_name* and *year*, and the rows in which *bpm* is greater than 120 and *trope_count* is greater than 5.

<!-- grade_id: select-answer, points: 0, grade: false, solution: true -->
```python
fast_corny = None

fast_corny = songs[(songs["bpm"] > 120) & (songs["trope_count"] > 5)][["song_name", "year"]]

fast_corny
```

Output:
```
song_name  year
school                          
Notre Dame   Victory March  1908
TCU         TCU Fight Song  1928
Illinois     Oskee-Wow-Wow  1910
Alabama        Yea Alabama  1926
Syracuse    Down the Field  1914
```

<!-- grade_id: select-tests, points: 1, grade: true, solution: false -->
```python
# TESTS
assert set(fast_corny.columns) == {"song_name", "year"}
assert len(fast_corny) == 5
print("OK")
```

Output:
```
OK
```

<!-- grade_id: select-grade, points: 3, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 8.
Assign to `short_conferences` a list or NumPy array containing the names of all conferences that have at least one song of duration less than 45 seconds.

<!-- grade_id: short-answer, points: 0, grade: false, solution: true -->
```python
short_conferences = None

short_songs = songs[songs["song_duration"] < 45]
short_conferences = short_songs["conference"].unique()

short_conferences = list(short_conferences)

print(short_conferences)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[35], line 2
      1 short_conferences = None
----> 2 short_conferences = df[df['sec_duration'] < 45]['conference'].unique()
      4 print(short_conferences)

NameError: name 'df' is not defined
```

<!-- grade_id: short-tests, points: 1, grade: true, solution: false -->
```python
# TESTS
assert type(short_conferences) in [np.ndarray, list]
assert len(short_conferences) == 4
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[36], line 2
      1 # TESTS
----> 2 assert type(short_conferences) in [np.ndarray, list]
      3 assert len(short_conferences) == 4

AssertionError:
```

<!-- grade_id: short-grade, points: 3, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

# 9.

Merge the `uni` frame with `songs`. Both should be merged on their index columns, and it should be an "inner" join (must be present in both to be included).

<!-- grade_id: merge-answer, points: 0, grade: false, solution: true -->
```python
combined = None
combined = uni.merge(songs, left_index=True, right_index=True, how='inner')
combined
```

Output:
```
Founded     Mascot conference            song_name  \
Rutgers     1766  Sir Henry    Big Ten  The Bells Must Ring   

                                 writers  year student_writer official_song  \
Rutgers  Richard Hadden and W.E. Sanford  1931            Yes           Yes   

        contest  bpm  ...  win_won victory_win_won  rah nonsense colors  men  \
Rutgers     Yes   75  ...      Yes             Yes  Yes      Yes    Yes  Yes   

        opponents spelling trope_count              spotify_id  
Rutgers        No      Yes           8  7EdGMXzZjJKNHJaszbWTHi  

[1 rows x 24 columns]
```

<!-- grade_id: merge-tests, points: 1, grade: true, solution: false -->
```python
# TESTS
assert type(combined) == pd.DataFrame
assert combined.shape == (1, 24)
print("OK")
```

Output:
```
OK
```

<!-- grade_id: merge-grade, points: 2, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```