```python
import pandas as pd
```

## 1.
Import the file `US_COVID-19_by_State_over_Time.csv` as a dataframe named `covid`. Display the first 8 rows.

<!-- grade_id: covid-display, points: 2, grade: true, solution: true -->
```python
import pandas as pd

covid = pd.read_csv('US_COVID-19_by_State_over_Time.csv')

covid.head(8)
```

Output:
```
submission_date state  tot_cases  conf_cases  prob_cases  new_case  \
0      03/11/2021    KS     297229    241035.0     56194.0         0   
1      12/01/2021    ND     163565    135705.0     27860.0       589   
2      01/02/2022    AS         11         NaN         NaN         0   
3      11/22/2021    AL     841461    620483.0    220978.0       703   
4      05/30/2022    AK     251425         NaN         NaN         0   
5      05/17/2020   RMI          0         0.0         0.0         0   
6      04/03/2020    ND        173         NaN         NaN        14   
7      09/04/2021    PR     173967    144788.0     29179.0       667   

   pnew_case  tot_death  conf_death  prob_death  new_death  pnew_death  \
0        0.0       4851         NaN         NaN          0         0.0   
1      220.0       1907         NaN         NaN          9         0.0   
2        0.0          0         NaN         NaN          0         0.0   
3      357.0      16377     12727.0      3650.0          
```

## 2.
Determine how many different states/regions are represented in the data.

<!-- grade_id: states-answer, points: 0, grade: false, solution: true -->
```python
num_states = None

num_states = covid['state'].nunique()

print(f"There are {num_states} unique states.")
```

Output:
```
There are 60 unique states.
```

<!-- grade_id: states-tests, points: 1, grade: true, solution: false -->
```python
# TESTS
assert type(num_states) == int
assert num_states > 50
```

<!-- grade_id: states-graded, points: 5, grade: true, solution: false -->
```python
# Intentionally blank--do not delete
```

## 3.
As you can see in the columns *conf_cases* and *prob_cases*, some of the reports do not include separate values for confirmed and probable cases and deaths. Remove all such rows from `covid`.

<!-- grade_id: drop-answer, points: 0, grade: false, solution: true -->
```python
covid_cleaned = covid.dropna(subset=['conf_cases', 'prob_cases'], how='any')
print(f"Cleaned length: {len(covid_cleaned)}")
covid = covid_cleaned
covid.head()
```

Output:
```
Cleaned length: 33962
  submission_date state  tot_cases  conf_cases  prob_cases  new_case  \
0      03/11/2021    KS     297229    241035.0     56194.0         0   
1      12/01/2021    ND     163565    135705.0     27860.0       589   
3      11/22/2021    AL     841461    620483.0    220978.0       703   
5      05/17/2020   RMI          0         0.0         0.0         0   
7      09/04/2021    PR     173967    144788.0     29179.0       667   

   pnew_case  tot_death  conf_death  prob_death  new_death  pnew_death  \
0        0.0       4851         NaN         NaN          0         0.0   
1      220.0       1907         NaN         NaN          9         0.0   
3      357.0      16377     12727.0      3650.0          7         3.0   
5        0.0          0         0.0         0.0          0         0.0   
7      274.0       2911      2482.0       429.0          8         3.0   

               created_at consent_cases consent_deaths  
0  03/12/2021 03:20:13 PM         Agree    
```

<!-- grade_id: drop-tests, points: 2, grade: true, solution: false -->
```python
# TESTS
assert len(covid) == 27977
assert covid["new_case"].sum() == 48938680
print("OK")
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[15], line 2
      1 # TESTS
----> 2 assert len(covid) == 27977
      3 assert covid["new_case"].sum() == 48938680
      4 print("OK")

AssertionError:
```

## 4.
Create a frame called `ratings` by loading the file `corporate_rating.csv`. Display the first 5 rows.

<!-- grade_id: ratings-display, points: 2, grade: true, solution: true -->
```python
ratings = None

ratings = pd.read_csv('corporate_rating.csv')

ratings.head()
```

Output:
```
Rating                   Name Symbol                  Rating Agency Name  \
0      A  Whirlpool Corporation    WHR          Egan-Jones Ratings Company   
1    BBB  Whirlpool Corporation    WHR          Egan-Jones Ratings Company   
2    BBB  Whirlpool Corporation    WHR                       Fitch Ratings   
3    BBB  Whirlpool Corporation    WHR                       Fitch Ratings   
4    BBB  Whirlpool Corporation    WHR  Standard & Poor's Ratings Services   

         Date             Sector  currentRatio  quickRatio  cashRatio  \
0  11/27/2015  Consumer Durables      0.945894    0.426395   0.099690   
1   2/13/2014  Consumer Durables      1.033559    0.498234   0.203120   
2    3/6/2015  Consumer Durables      0.963703    0.451505   0.122099   
3   6/15/2012  Consumer Durables      1.019851    0.510402   0.176116   
4  10/24/2016  Consumer Durables      0.957844    0.495432   0.141608   

   daysOfSalesOutstanding  ...  effectiveTaxRate  \
0               44.203245  ...          0.
```

## 5.
The ratings are ordered AAA, AA, A, BBB, BB, B, CCC, CC, C, D. Create a new column called *Rating_number* in which each string value in *Rating* column is replaced with the ordinal equivalents 1, 2, 3, ..., 10.

<!-- grade_id: replace-answer, points: 0, grade: false, solution: true -->
```python
# Change the frame called ratings

rating_map = {
    'AAA': 1, 'AA': 2, 'A': 3, 'BBB': 4, 'BB': 5, 'B': 6, 
    'CCC': 7, 'CC': 8, 'C': 9, 'D': 10
}

ratings['Rating_number'] = ratings['Rating'].map(rating_map)

ratings[["Name", "Rating", "Rating_number"]].head(8)
```

Output:
```
Name Rating  Rating_number
0  Whirlpool Corporation      A              3
1  Whirlpool Corporation    BBB              4
2  Whirlpool Corporation    BBB              4
3  Whirlpool Corporation    BBB              4
4  Whirlpool Corporation    BBB              4
5      Schlumberger N.V.     AA              2
6      Schlumberger N.V.     AA              2
7      Schlumberger N.V.      A              3
```

<!-- grade_id: replace-tests, points: 4, grade: true, solution: false -->
```python
# TESTS
assert len(ratings.columns) == 32
assert ratings["Rating_number"].dtype in ["int64", "int"]
assert ratings["Rating_number"].min() == 1
assert ratings["Rating_number"].max() == 10
assert ratings["Rating_number"].sum() == 8841
print("OK")
```

Output:
```
OK
```